Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Munachimso Aneke"
COLLABORATORS = ""
```

---

# Homework 4

## Question 1

Create a dictionary in Python called `lotr` that lists: `'race': <count of characters with that race>` of the following characters in "Lord of the Rings:" <br>

- Legolas (elf)
- Haldir (elf)
- Galadriel (elf)
- Aragorn (human)
- Boromir (human)
- Faramir (human)
- Eomer (human)
- Eowyn (human)
- Frodo (hobbit)
- Sam (hobbit)
- Merry (hobbit)
- Pippin (hobbit

<!-- grade_id: cell-b5a249261ec2933e, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
lotr = {'elf': 3, 'human':5, 'hobbit': 4}
lotr
```

Output:
```
{'elf': 3, 'human': 5, 'hobbit': 4}
```

<!-- grade_id: cell-2506f5830af69406, points: 2, grade: true, solution: false -->
```python
## question 1 - check whether lotr was created and is a dict type
assert type(lotr) == dict
assert lotr['elf'] == 3
```

<!-- grade_id: cell-0ce3e5d53d11cdaf, points: 2, grade: true, solution: false -->
```python

```

## Question 2
Create a function `intgr(n)` that returns if an inputted value `n` is an integer or not.  <br>

For example: <br>
if you input `intgr(5.6)` the output should be `non-integer` <br>
if you input `intgr(6)` the output should be `integer`.

<!-- grade_id: cell-edb5ab58f3912501, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
def intgr(n):
    if isinstance(n,int):
        c = "integer"
    else: 
        c = "non-integer"
    return(c)

intgr(5.6)
#intgr(6)
```

<!-- grade_id: cell-e674f5ea70fb7460, points: 2, grade: true, solution: false -->
```python
## question 2 - check whether intgr was created and outputs match expectations
assert type(intgr(6)) == str
```

<!-- grade_id: cell-f75b32e0c8a7f502, points: 2, grade: true, solution: false -->
```python

```

## Question 3
Write a function `mirror(s)` that returns a mirrored version of the string (`s`) (i.e., the string concatenated with the string in reverse).
<br>
Examples: <br>
`mirror('1')` returns `'11'` <br>
`mirror('RaInBoW')` returns `'RaInBoWWoBnIaR'` <br>
`mirror('Aa')` returns `'AaaA'` <br>
`mirror('orange')` returns `'orangeegnaro'` <br>

<!-- grade_id: cell-95d76d9e4df725db, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
def mirror(s):
    c = s[:]
    for i in range(len(s)):
        c = c+s[len(s)-1-i]
    return(c)

mirror('RaInBoW')
```

Output:
```
'RaInBoWWoBnIaR'
```

<!-- grade_id: cell-e3be9874e269b4c3, points: 2, grade: true, solution: false -->
```python
## question 3 - check whether mirror was created and if mirror(1) returns str type
assert type(mirror('1')) == str
```

<!-- grade_id: cell-f4afb69a1ebde960, points: 2, grade: true, solution: false -->
```python

```

## Question 4
Create a function, `numlist(x, y)`, that returns a list of all of the integers between integer `x` and integer `y` inclusive.  The function should return `number 1 must be smaller than number 2` if `x` > `y`.  The function should return `x and y are both not integers` if either `x` or `y` are not integers.

For example: <br>
`numlist(2, 5)` will return `[2, 3, 4, 5]` <br>
`numlist(5, 2)` will return `number 1 must be smaller than number 2` <br>
`numlist(2.3, 4)` will return `x and y are both not integers`

<!-- grade_id: cell-0d1f2ae2d9989d25, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
def numlist(x,y):
    if type(x) != int or type(y) != int:
        c = "x and y are both not integers"
    elif x > y:
        c = "number 1 must be smaller than number 2"
    else:
        c = str(list(range(x,y+1)))
    return(c)

numlist(2, 8)
```

Output:
```
'[2]'
```

<!-- grade_id: cell-7d3dd358972d8e2b, points: 2, grade: true, solution: false -->
```python
## Question 4 - check that numlist was created and that numlist(2.3, 4) return a str type
assert type(numlist(2.3, 4)) == str
```

<!-- grade_id: cell-4c3520c383b0f7e3, points: 2, grade: true, solution: false -->
```python

```

## Question 5
Write a function `name(s)` that takes in a full name `s` formatted as the string `Last, First` and output the name formatted as the string `First Last`.  If the inputted string does not contain a comma, return the inputted string. If the inputted string does not contain a space, add a space in the output. <br>
<br>
For example: <br>
`name('Smith, Joe')` returns `'Joe Smith'` <br>
`name('Arnold, Benedict')` returns `'Benedict Arnold'` <br>
`name('Johnson Jr., Sara')` returns `'Sara Johnson Jr.'` <br>
`name('Gerald')` returns `'Gerald'` <br>
`name('Griffin,Stephanie')` returns `'Stephanie Griffin'`

```python
s1 = 'Smith, Joe'
s2 = 'Johnson Jr., Sara'
s3 = 'Gerald'
s4 = 'Griffin,Stephanie'
s5 = 'Arnold, Benedict'

if ',' in s1:
    print("True")

if ',' in s2:
    print("True")

if ',' in s3:
    print("True")
    

c = s.split(',')
d = c[-1] + " " + c[0]
d
```

Output:
```
True
True
' Joe Smith'
```

<!-- grade_id: cell-0a718c50cc698675, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
def name(s):
    if ',' in s: # commas = true
        c1 = s.split(',')

        if ' ' in c1[1]: # if there's a space before the first name
            c2 = c1[1].split(' ')[1] # take the first name
        else:
            c2 = c1[1][:]

        c = c2 + " " + c1[0]
    else:  # no commas
        c = s
    return(c)

print(name('Smith, Joe'))
print(name('Johnson Jr., Sara'))
print(name('Gerald'))
print(name('Griffin,Stephanie'))
```

Output:
```
Joe Smith
Sara Johnson Jr.
Gerald
Stephanie Griffin
```

<!-- grade_id: cell-03ef629409933723, points: 2, grade: true, solution: false -->
```python
## Question 5 - check to see if name was created and returns a str type
assert type(name('Arold, Benedict')) == str
```

<!-- grade_id: cell-bc2692e7579c5798, points: 3, grade: true, solution: false -->
```python

```

## Question 6
Write a function `concat(s1, s2)` that contactonates two strings with the longer string first.
If `s1` and `s2` strings are the same length `concat` will return `s1s2`.

<br>
For example: <br>
`concat('cat', 'dog')` returns `catdog` <br>
`concat('cat', 'dogs')` returns `dogscat` <br>
`concat('123', '45')` returns `12345` <br>

<!-- grade_id: cell-e2146eb089f01bed, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError

# longer version of concat()
"""
def concat(s1, s2):
    if len(s1) > len(s2):
        c = s1 + s2
    elif len(s1) < len(s2):
        c = s2 + s1
    else:
        c = s1 + s2
    return(c)
"""

# shorter version of concat()
def concat(s1, s2):
    if len(s1) < len(s2):
        c = s2 + s1
    else: # if len(s1) >= len(s2)
        c = s1 + s2
    return(c)

print(concat('cat', 'dog'))
print(concat('cat', 'dogs'))
print(concat('123', '45'))
```

Output:
```
catdog
dogscat
12345
'34'
```

<!-- grade_id: cell-7ea4474f90cfc9c3, points: 2, grade: true, solution: false -->
```python
## Question 6 - check to see if concat was created and returns a str
assert type(concat('3', '4')) == str
```

<!-- grade_id: cell-dd5fd7f43ed0f332, points: 3, grade: true, solution: false -->
```python

```

## Question 7
Write a function `chrctr(s, n, c)` that returns "True" if the `n`th character in the string `s` is equal to `c`.  If `n` is a number larger than the length of the character string `s`, return "False".

Parameters: <br>
- `s` is a character string <br>
- `n` is an integer <br>
- `c` is a character string <br>

For example: <br>
`chrctr('cat', 2, 'a')` returns `True` <br>
`chrctr('1234', 4, '5')` returns `False` <br>
`chrctr('1234', 7, '4')` returns `False` <br>

<!-- grade_id: cell-dfcf64ae9655fcd6, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError

def chrctr(s, n, c):
    if n> len(s):
        return False
    else:
        if s[n-1] == c:
            return True
        else:
            return False

chrctr('cat', 2, 'a')
#chrctr('1234', 4, '5')
#chrctr('1234', 7, '4')
```

Output:
```
True
```

<!-- grade_id: cell-7bd6b124e1b6be54, points: 2, grade: true, solution: false -->
```python
## Question 7 - check to see if chrctr was created and returns bool type
assert type(chrctr('cat', 2, 'a')) == bool
```

<!-- grade_id: cell-087b489860322cff, points: 3, grade: true, solution: false -->
```python

```

## Question 8A

Set `bln` equal to a number that when you cast `bln` as a Boolean the response would be `False`.

<!-- grade_id: cell-3d2e1ea5bf197f8c, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
bln = 0
bool(bln)
```

Output:
```
False
```

<!-- grade_id: cell-312ff1c6dae63fc8, points: 2, grade: true, solution: false -->
```python
## Question 8a - check to see if bln was created and is an int value
assert type(bln) == int
```

<!-- grade_id: cell-6d7d0579b5dc4543, points: 2, grade: true, solution: false -->
```python

```

## Question 8B
Create a new list, `L3` that is the following two lists concatonated: <br>
  L1 = [2, 4, 6, 8] <br>
  L2 = ["Star", "Hexagon"] <br>
And then append the name "Asterisk" between "Star" and "Hexagon" in the list.

<!-- grade_id: cell-c3acd3eee7a6657e, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
L1 = [2, 4, 6, 8]
L2 = ["Star", "Hexagon"]
L3 = L1 + L2
L3.insert(5, "Asterisk")
L3
```

Output:
```
[2, 4, 6, 8, 'Star', 'Asterisk', 'Hexagon']
```

<!-- grade_id: cell-0cd4a639ea1c989a, points: 2, grade: true, solution: false -->
```python
## Question 8b - check to see if L3 was created and is length 7
assert len(L3) == 7
```

<!-- grade_id: cell-fd9afe22ac7d9647, points: 2, grade: true, solution: false -->
```python

```

## Question 9

Create a function called `greet(a, b)`.  For two variables, `a` and `b`, return "Hello World" if `a` is greater than `b` else return "Goodbye World" .

<!-- grade_id: cell-8abe3634d2c227d3, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
def greet(a,b):
    if a > b:
        c = "Hello World"
    else:
        c = "Goodbye World"
    return(c)

# greet(2, 3)
greet(5, 3)
```

Output:
```
'Hello World'
```

<!-- grade_id: cell-17bcdec4bb80055c, points: 2, grade: true, solution: false -->
```python
## Question 9 - check to see if greet was created and is type str
assert type(greet(2, 3)) == str
```

<!-- grade_id: cell-f10ba0e12e666066, points: 2, grade: true, solution: false -->
```python

```

## Question 10
Create a function `funr(rang, num)` that will return 0 if the number `num` is not in the range `rang` and 1 if the number `num` is in the range `rang`.
<br>
Ex: <br>
`rang` = `range(2, 20, 3)`, `num` = `4`, `funr(rang, num)` = `0` <br>
`rang` = `range(2, 20, 3)`, `num` = `5`, `funr(rang, num)` = `1`

<!-- grade_id: cell-d0baf5b61f6fbe0e, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError

def funr(rang, num):
    if num in list(rang):
        c = 1
    else:
        c = 0
    return(c)

rang = range(2, 20, 3)
#num = 5
num = 4
funr(rang, num)
```

Output:
```
1
```

<!-- grade_id: cell-63900600ed3dcfe7, points: 2, grade: true, solution: false -->
```python
## Question 10 - check to see if funr was created and funr() returns int

rang = range(2, 20, 3)
num = 4

type(funr(rang, num)) == int
```

Output:
```
True
```

<!-- grade_id: cell-ae696809f9d75a80, points: 2, grade: true, solution: false -->
```python

```

## Question 11

Which keyword indicates an anonymous function? <br>

a. None <br>
b. anon <br>
c. global <br>
d. lambda <br>
e. theta <br>

Answer by naming `ans = " "` with correct letter in quotes.  `ans` should be a `str` type.

<!-- grade_id: cell-b4265d3f154883ab, points: 0, grade: false, solution: true -->
```python
# your code here
# raise NotImplementedError
ans = "d"
```

<!-- grade_id: cell-d7146f0b580ad9d1, points: 1, grade: true, solution: false -->
```python
## Question 11 - check to see if ans was created and is type str
assert type(ans) == str
```

<!-- grade_id: cell-119ec2d9e68176e0, points: 2, grade: true, solution: false -->
```python

```